# Lab 1: Getting started 

The goals of this lab are:
1. Probability warmup
2. Familiarize ourselves with Jupyter notebooks
3. Use Python in the Jupyter Notebook
4. Perform mathmatical computations using the NumPy library
5. Import and manipulate data using the Pandas library
6. Plot data with Matplotlib and Seaborn
7. Practice exploring new datasets and visualizing them

## Probability Practice (Show all work)

### Question 1 (2 pts) 

My neighbor has two children. Assuming that the gender of a child is like a coin flip, it is most likely, a priori, that my neighbor has one boy and one girl, with probability $1/2$. The other possibilities—two boys or two girls—have probabilities $1/4$ and $1/4$.

a. Suppose I ask him whether he has any boys, and he says yes. What is the probability that one child is a girl?

b. Suppose instead that I happen to see one of his children run by, and it is a boy. What is the probability that the other child is a girl?


**Your solution:**

a. 2/3, we have confirmed that 1 child is a boy, leaving 3 cases. 2 cases of 1 boy, 1 girl, or 1 case of 2 boys. 2 cases have 1 girl out of 3 possible cases.

b. 1/2, (2 * (1 girl, 1 boy case) + 2 * (0 girl case, either boy may have been seen))/(4 possible cases) = 2/4 = 1/2

### Question 2 (2 pts)

In 1995, they introduced blue M&M’s. Before then, the color mix in a bag of plain M&M’s was $30\%$ Brown, $20\%$ Yellow,
$20\%$ Red, $10\%$ Green, $10\%$ Orange, $10\%$ Tan. Afterward it was $24\%$ Blue, $20\%$ Green, $16\%$ Orange, $14\%$ Yellow, 
$13\%$ Red, $13\%$ Brown.

Suppose a friend of mine has two bags of M&M’s, and he tells me that one is from 1994 and one from 1996. 
He won’t tell me which is which, but he gives me one M&M from each bag. 
One is yellow and one is green. What is the probability that the yellow one came from the 1994 bag?


**Your solution:**

(0.2/0.34 * 0.2/0.3)/(0.2/0.34 * 0.2/0.3 + 0.14/0.34 * 0.1/0.3) = 0.74, about 74%

## Jupyter notebooks

Jupyter notebooks are useful for tinkering because they're very readable, but research projects that require more abstraction (more lines of code, or cyclomatic complexity, or Halstead complexity, or etc.) should be coded in Python scripts using an object-oriented programming paradigm in an integrated development environment with a more full suite of features. For this course, the Lab$n$_student.ipynb notebooks provided are sufficient.

Via a JupyterHub hosted by UCLA's IDRE, your notebooks will be in your own Docker container, where all the libraries you will need for the assignments should be installed. We recommend Chrome for accessing JupyterHub.

Tips: 
- Save your notebook frequently (CMD+S on Mac.)
- Be careful about the order in which you run cells.

### Resources:
- Assignment notebooks will contain some useful links.
- Stackoverflow is a great resource for troubleshooting.
- [This book](https://jakevdp.github.io/PythonDataScienceHandbook/) contains many good introductions to topics we'll discuss in the class, and it includes Python code written in Jupyter Notebooks.
- Here is a [Jupyter Notebook cheatsheet](https://www.edureka.co/blog/wp-content/uploads/2018/10/Jupyter_Notebook_CheatSheet_Edureka.pdf).

### Two Useful Commands
The code in the cell below imports a package.  You can run the cell by clicking the button at the top of this webpage, or using the shortcuts CTRL/CMD+ENTER or SHIFT+ENTER to run the cell and keep the cell selected or to run the cell then select the next cell, respectively.  Also, here is a link to the [NumPy documentation](https://numpy.org/doc/).

```python
from numpy.testing import assert_array_almost_equal, assert_almost_equal, assert_array_equal
import numpy as np
```

Using an exclamation point, one can run shell commands inside of a Jupyter Notebook cell.  For example, the following cell issues the ``pwd`` command, which prints the working directory.

```python
! pwd
```

Output:
```
'pwd' is not recognized as an internal or external command,
operable program or batch file.
```

### Question 3 (1 pt)
Run the ``%magic`` command, then try out one of the commands listed in the resulting display. Make sure to delete the output cell before submitting the lab due to the significant length.

<!-- grade_id: magics2, points: 1, grade: true, solution: true -->
**Your solution:**

```python
%magic
```

Output:
```
IPython's 'magic' functions
===========================

The magic function system provides a series of functions which allow you to
control the behavior of IPython itself, plus a lot of system-type
features. There are two kinds of magics, line-oriented and cell-oriented.

Line magics are prefixed with the % character and work much like OS
command-line calls: they get as an argument the rest of the line, where
arguments are passed without parentheses or quotes.  For example, this will
time the given statement::

        %timeit range(1000)

Cell magics are prefixed with a double %%, and they are functions that get as
an argument not only the rest of the line, but also the lines below it in a
separate argument.  These magics are called with two arguments: the rest of the
call line and the body of the cell, consisting of the lines below the first.
For example::

        %%timeit x = numpy.random.randn((100, 100))
        numpy.linalg.svd(x)

will time the execution of the numpy svd routi
```

```python
%pylab
```

Output:
```
Using matplotlib backend: <object object at 0x0000011DF527B710>
%pylab is deprecated, use %matplotlib inline and import the required libraries.
Populating the interactive namespace from numpy and matplotlib
```

### Question 4 (1 pt)

It's useful to know where python is located on your computer and what version of python is installed; try out the following terminal commands: ``! which python`` and ``! python --version``

<!-- grade_id: magics2, points: 1, grade: true, solution: true -->
**Your solution:**

```python
! which python
```

Output:
```
'which' is not recognized as an internal or external command,
operable program or batch file.
```

```python
! python --version
```

Output:
```
Python 3.11.5
```

### Question 5 (1 pt)

Jupyter also supports LaTeX in markdown (double-click this cell to see the markdown syntax). For inline math, surround expressions with single dollar signs like this: \\$e^{i\pi} + 1 = 0\\$  $ \rightarrow e^{i\pi} + 1 = 0$

Double dollar signs around expressions produce equations: \\$\\$ e^x=\sum_{i=0}^\infty \frac{1}{i!}x^i \\$\\$ becomes $$e^x=\sum_{i=0}^\infty \frac{1}{i!}x^i$$


<b style="color:red">Type out any LaTeX equation to test it out in the cell below:</b>

$\sin(2x) = 2\cos(x)\sin(x)$

## Python Review

### Question 6 (2 pts)

First, define a function called ``mult`` that accepts two variables, multiplies the variables, and returns the product.

<!-- grade_id: mult_function, points: 1, grade: true, solution: true -->
**Your solution:**

```python
def mult(a, b):
    return a * b
```

<!-- grade_id: call_mult, points: 0, grade: false, solution: true -->
Next, call ``mult`` and pass it 5 and 3. Assign the output to a variable named ``product``.

**Your solution:**

```python
product = mult(5,3)
```

Modify ``mult`` so that, instead of returning the product of two integer or float inputs, it returns the dot product between the vectors defined by two lists or NumPy arrays (*hint:* look at the [documentation for the linear algebra methods in NumPy](https://numpy.org/doc/stable/reference/routines.linalg.html).)  With your new function, assign the output of the product between arrays ``[1,2,3]`` and ``[4,5,6]`` to a variable valled ``dotprod``, then print ``dotprod``.

<!-- grade_id: multarr, points: 0, grade: false, solution: true -->
**Your solution:**

```python
def mult(arr1, arr2):
    if (not isinstance(arr1, list) or not isinstance(arr2, list) or len(arr1) != len(arr2)):
        raise ArithmeticError("Either not a vector or the two vectors do not have the same size ")
        return -1
    res = 0
    for x in range(0, len(arr1)):
        res += arr1[x] * arr2[x]
    return res
```

<!-- grade_id: dotprod, points: 1, grade: true, solution: false -->
```python
mult([1,2,3],[4,5,6])
```

Output:
```
32
```

### Lists and Arrays
Array algebra is the core of lots of machine learning, and Numpy is the package we'll use to store and operate on arrays. *Note* that one important challenge is minimizing the computational cost of array algebra.

### Question 7 (2 pts)

In the cell below, assign a list containing integers spanning the interval between 0 and 9 to a variable named ``l``.  Use a NumPy method to generate the list; don't type in each integer. *Note* that, if the method you use generates an array, you'll need to convert the array to a list.

<!-- grade_id: list, points: 0, grade: false, solution: true -->
**Your solution:**

```python
l = list(np.arange(0, 10))
```

Run the cell below to verify that ``l`` is the desired list.

<!-- grade_id: list_test, points: 1, grade: true, solution: false -->
```python
assert l == [0,1,2,3,4,5,6,7,8,9]
```

Using [list comprehension](https://docs.python.org/3/tutorial/datastructures.html#list-comprehensions), generate the same list and assign it to a variable called ``l2``.

**Your solution:**

<!-- grade_id: list_comprehension, points: 0, grade: false, solution: true -->
```python
l2 = list(map(lambda x: x, range(10)))
```

Run the cell below to verify that ``l2`` is the desired list.

<!-- grade_id: list_comp, points: 1, grade: true, solution: false -->
```python
assert l2 == [0,1,2,3,4,5,6,7,8,9]
```

Using the ``np.array()`` [method](https://numpy.org/doc/stable/reference/generated/numpy.array.html), assign an array with the same elements as the list ``l2`` to a variable named ``arr`` (generate the array using ``l2``.)

<!-- grade_id: arr_of_floats, points: 0, grade: false, solution: true -->
**Your solution:**

```python
arr = np.array(l2)
```

Run the cell below to verify.

<!-- grade_id: arr_test, points: 1, grade: true, solution: false -->
```python
assert isinstance(arr, np.ndarray)
```

### Question 8 (7 pts)

Unlike lists, arrays can have multiple dimensions of data.  Non-vector tensors are the bulk of the arrays which appear in machine learning problems.  Using a [NumPy method](https://numpy.org/doc/stable/reference/generated/numpy.ones.html), create a $3\times5$ array called ``arr2`` with $1$ for every element.

<!-- grade_id: arr_init_ones, points: 0, grade: false, solution: true -->
**Your solution:**

```python
arr2 = np.ones((3, 5))
```

Verify your solution:

<!-- grade_id: arr_init_ones_test, points: 1, grade: true, solution: false -->
```python
assert_array_equal(arr2,np.ones((3,5)))
```

Another common task in ML is to create arrays filled with random numbers.  There are lots of different NumPy methods for generating random numbers!  Create a $3 \times 3$ array of random integers normally distributed on the interval ``[0,10]``, and assign it to a variable called ``arr3``.

<!-- grade_id: rand_arr, points: 0, grade: false, solution: true -->
**Your solution:**

```python
arr3 = np.random.rand(3, 3)
arr3 *= 10
```

Verify the shape of your array:

<!-- grade_id: rand_arr_test, points: 1, grade: true, solution: false -->
```python
assert arr3.shape == (3,3)
```

Assign the ``shape`` attribute of ``arr3`` to a variable called ``s``.

<!-- grade_id: arr3, points: 0, grade: false, solution: true -->
**Your solution:**

```python
s = arr3.shape
```

Verify the type of ``s``:

<!-- grade_id: arr3_tests, points: 1, grade: true, solution: false -->
```python
assert isinstance(s, tuple)
```

Assign the number of array dimensions of ``arr3`` to a variable called ``d``.  You can find the right method in this [guide](https://note.nkmk.me/en/python-numpy-ndarray-ndim-shape-size/).

**Your solution:**

<!-- grade_id: arr3_cont, points: 0, grade: false, solution: true -->
```python
d = ndim(arr3)
```

<!-- grade_id: arr3_cont2, points: 0, grade: false, solution: true -->
Using [array indexing](https://numpy.org/doc/stable/user/basics.indexing.html), print the element of ``arr3`` in the first row and first column and assign it to a variable called ``x``.

**Your solution:**

```python
x = arr3[0][0]
```

<!-- grade_id: arr3_cont2_test, points: 1, grade: true, solution: false -->
```python
assert x == arr3[0,0]
```

<!-- grade_id: lasts, points: 0, grade: false, solution: true -->
Create a variable ``e`` with the value of the element in the row with the largest index and column with the largest index in ``arr3``.  *Note* that there is a way to index the last column or row that doesn't depend on the size of the array.

**Your solution:**

```python
e = arr3[-1][-1]
```

```python
assert e == arr3[-1,-1]
```

Since datasets are often large in ML contexts, we need to slice arrays and only use a portion of a dataset.  One way to slice an array is using the command ``array[start_index:stop_index]``.  We can also specify a step size, ``array[start_index:stop_index:step_size]``.  Create a new variable called ``y`` which contains just the first two rows of ``arr3``.

**Your solution:**

<!-- grade_id: arr_slice, points: 0, grade: false, solution: true -->
```python
y = arr3[:2]
```

<!-- grade_id: slice_test, points: 1, grade: true, solution: false -->
```python
assert y.all() == arr3[:2].all()
```

Using NumPy's ``arange`` and ``reshape`` methods, create an array called ``arr55`` so that ``print(arr55)`` returns this:

``[[1,2,3,4,5],
  [6,7,8,9,10],
  [11,12,13,14,15],
  [16,17,18,19,20],
  [21,22,23,24,25]]``

<!-- grade_id: array_shape, points: 0, grade: false, solution: true -->
**Your solution:**

```python
arr55 = np.arange(1, 26).reshape(5, 5)
```

<!-- grade_id: array_shape_tests, points: 1, grade: true, solution: false -->
```python
assert arr55.all() == np.arange(1,26).reshape(5,5).all()
```

## Broadcasting: 

Broadcasting, defined by VanderPlas: a set of rules for applying binary ufuncs (e.g., addition, multiplication, etc.)
on arrays of different sizes. For arrays of the same size, binary operations are performed on an
element-by-element basis. Broadcasting allows these operations to be performed on arrays of different
sizes, for example adding a scalar to an array. Here's a brief demonstration of how it works:

**Adding two arrays with different shapes:**

First, define an array with shape (4,1).

``a = np.array(
        [[1.],
         [2.],
         [3.],
         [4.]])``

Note that the periods after the numbers tell Python that the elements should have *float* DataTypes rather than *integer*. Now define another array with shape (3,).

``b = np.array([0., 1., 2.])``

Now we operate on the arrays:

``c = a + b``

And find that

``c = np.array([[1., 2., 3.],
       [2., 3., 4.],
       [3., 4., 5.],
       [4., 5., 6.]])``

Here's what happened inside the NumPy machinery:

NumPy aligned the shapes of a and b on the last axis and prepended 1s to the shape with fewer axes:
 a: 4 x 1     --->    a: 4 x 1
 b:     3     --->    b: 1 x 3
Then checked that the sizes of the axes matched or were equal to 1:  
 a: 4 x 1  
 b: 1 x 3  
``a`` and ``b`` satisfied this criterion.  
NumPy stretched both arrays on their 1-valued axes so that their shapes matched, then added them together.
``a`` was replicated 3 times in the second axis, while ``b`` was replicated 4 times in the first axis.
This meant that the addition in the final step was  

``
[[1., 1., 1.],            [[0., 1., 2.],  
 [2., 2., 2.],      +      [0., 1., 2.],  
 [3., 3., 3.],             [0., 1., 2.],  
 [4., 4., 4.]]             [0., 1., 2.]]
 `` 

Addition was then carried out element-by-element, as you can verify by referring back to the output of the code cell above.
This resulted in an output with shape 4 x 3.

## Numpy's broadcasting rule
Broadcasting rules describe how values should be transmitted when the inputs to an operation do not match.
In numpy, the broadcasting rule is very simple:

Prepend 1s to the smaller shape,
check that the axes of both arrays have sizes that are equal or 1,
then stretch the arrays in their size-1 axes.

A crucial aspect of this rule is that it does not require the input arrays have the same number of axes.
Another consequence of it is that a broadcasting output will have the largest size of its inputs in each axis.
Take the following multiplication as an example:

    a: 3 x 7 x 1  
    b:     1 x 5  
a * b: 3 x 7 x 5
You can see that the output shape is the maximum of the sizes in each axis.

NumPy's broadcasting rule also does not require that one of the arrays has to be bigger in all axes.  
      
 
Broadcasting behavior also points to an efficient way to compute an outer product in NumPy:

```python
a = np.array([-1., 0., 1.])
b = np.array([0., 1., 2., 3.])

a[:, np.newaxis] * b  # outer product ab^T, where a and b are column vectors
```

Output:
```
array([[-0., -1., -2., -3.],
       [ 0.,  0.,  0.,  0.],
       [ 0.,  1.,  2.,  3.]])
```

The idea of numpy stretching the arrays in their size-1 axes is useful and is functionally correct. But this is not what numpy literally does behind the scenes, since that would be an inefficient use of memory. Instead, numpy carries out the operation by looping over singleton (size-1) dimensions.



### Question 9 (2 pts)
To give you some practise with broadcasting, <b style="color:red">try predicting the output shapes for the following operations</b>:

```python
# Run this cell, which defines three arrays with different shapes
a = np.array([[1.], [2.], [3.]])
b = np.zeros(shape=[10, 1, 1])
c = np.ones(shape=[4])
d = np.array([1,2,3,4])
```

What are the shapes of ``a``, ``b``, ``c``, and ``d``? Define ``d1`` as the shape of ``d``.

<!-- grade_id: tuple, points: 0, grade: false, solution: true -->
```python
d1 = np.shape(d)
```

<!-- grade_id: tuple_test, points: 1, grade: true, solution: false -->
```python
assert d1==(4,)
```

Multiply two arrays with differen shapes:

<!-- grade_id: cell-ba93b508e3a31c1e, points: 0, grade: false, solution: true -->
```python
a = np.array([[[0.01], [0.1]],
              [[1.00], [10.]]])  # what shape is this? it's (2, 2, 1)
b = np.array([[[2., 2.]],
              [[3., 3.]]])       # what shape is this? it's (2, 1, 2)

c = a * b 
```

What is the shape of ``c``? Type the answer below and assign to ``v``.

```python
v = c.shape
```

<!-- grade_id: array_mult, points: 1, grade: true, solution: false -->
```python
assert v == (2,2,2)
```

## Pandas:
Pandas is a Python library built on Numpy. It defines Series and  DataFrame objects, which are multi-dimensional arrays with row and column labels attached naturally. They are good with various data types and missing data. This is a link to the [Pandas documentation](https://pandas.pydata.org/docs/).

```python
import pandas as pd
```

A Pandas Series object is a one dimensional array of indexed data. Such an object can be created from a list in a number of ways, and so can DataFrames.  We will try to introduce the various ways; they'll be confusing for now, but it gets easier with more exposure; you will start to see the different, equivalent ways of performing ML and data science tasks.

### Question 10 (1 pt)

```python
data = pd.Series([0.30, 0.5, 0.78, 1.0])
```

In the cells below, try data.values and data.index, and try to index the Series using ``data[2]``.

<!-- grade_id: series1, points: 1, grade: true, solution: true -->
```python
data.values
```

Output:
```
array([0.3 , 0.5 , 0.78, 1.  ])
```

<!-- grade_id: series2, points: 1, grade: true, solution: true -->
```python
data.index
```

Output:
```
RangeIndex(start=0, stop=4, step=1)
```

```python
print(data[2])
```

Output:
```
0.78
```

### DataFrames

### Question 11 (8 pts)

You can think of a Pandas DataFrame object as a sequence of aligned Series objects (aligned in the sense that they share the same index.) These are some ways you can construct DataFrames: from a single Series object, from a list of dictionaries, from a dictionary of Series objects, from a 2-d numpy array, from a Numpy Structured array, or by importing some data (for example, a CSV file). If you know SQL, DataFrames are designed to allow for different joins of datasets, and lots of other functionalities.

Next, we are going to load a dataset and perform Pandas operations on the dataset.

We'll learn how to calculate statstics on datasets and how to renaming columns, add rows, drop columns, combine columns, create new columns out of existing ones, locate information with queries, filter, etc. 

<!-- grade_id: imports3, points: 0, grade: false, solution: true -->
```python
# (1 pt.)

import pandas as pd
import numpy as np
from sklearn.utils import Bunch
data_url = "http://lib.stat.cmu.edu/datasets/boston"
raw_df = pd.read_csv(data_url, sep=r"\s+", skiprows=22, header=None)
data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
target = raw_df.values[1::2, 2]
feature_names = np.array(['CRIM', 'ZN', 'INDUS', 'CHAS', 'NOX', 'RM', 'AGE', 'DIS', 'RAD',
        'TAX', 'PTRATIO', 'B', 'LSTAT'])
boston = Bunch(data=data, target=target, feature_names=feature_names)

# in the space below, print the datatype of boston.
### BEGIN SOLUTION
print(type(boston))
### END SOLUTION
```

Output:
```
<class 'sklearn.utils._bunch.Bunch'>
```

Print the datatype of ``boston`` in the cell below.

```python
print(type(boston))
```

Output:
```
<class 'sklearn.utils._bunch.Bunch'>
```

<!-- grade_id: boston_datatype, points: 1, grade: true, solution: false -->
```python
import sklearn as sk
assert isinstance(boston, sk.utils.Bunch) #bunch file has dicts of labels, separate target, etc.
```

(1pt.) 
<b style="color:red">
What is a bunch object? 
What is it made up of? Answer in the next cell.
</b>

A bunch object is a dictionary style object that maps keys to resulting values.

(1pt.) Transform just boston.data into a Pandas DataFrame (not the other items in the bunch) and name the dataframe ``Xbos``. *Hint 1*: you want to assign the columns to be boston.feature_names.  *Hint 2*: don't worry about the y value (the target,) we won't be using that for this exercise.

<!-- grade_id: df, points: 0, grade: false, solution: true -->
```python
Xbos = pd.DataFrame(boston.data, columns = boston.feature_names)
```

<!-- grade_id: df_test, points: 1, grade: true, solution: false -->
```python
from pandas.testing import assert_frame_equal
assert_frame_equal (Xbos, pd.DataFrame(boston.data, columns = boston.feature_names))
```

(1pt.) 
Check out your dataframe! in the space below, try ``Xbos.head()``, ``Xbos['insert_col_name'].value_counts()``,
``Xbos.values``, ``Xbos.std()``, and ``Xbos.mean()``. Refer back to the documentation in sklearn to understand the column values.

```python
Xbos.head()
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM   AGE     DIS  RAD    TAX  \
0  0.00632  18.0   2.31   0.0  0.538  6.575  65.2  4.0900  1.0  296.0   
1  0.02731   0.0   7.07   0.0  0.469  6.421  78.9  4.9671  2.0  242.0   
2  0.02729   0.0   7.07   0.0  0.469  7.185  61.1  4.9671  2.0  242.0   
3  0.03237   0.0   2.18   0.0  0.458  6.998  45.8  6.0622  3.0  222.0   
4  0.06905   0.0   2.18   0.0  0.458  7.147  54.2  6.0622  3.0  222.0   

   PTRATIO       B  LSTAT  
0     15.3  396.90   4.98  
1     17.8  396.90   9.14  
2     17.8  392.83   4.03  
3     18.7  394.63   2.94  
4     18.7  396.90   5.33
```

```python
Xbos['CRIM'].value_counts()
```

Output:
```
CRIM
0.01501     2
14.33370    2
0.03466     1
0.03113     1
0.03049     1
           ..
1.51902     1
1.83377     1
1.46336     1
1.27346     1
0.04741     1
Name: count, Length: 504, dtype: int64
```

```python
Xbos.values
```

Output:
```
array([[6.3200e-03, 1.8000e+01, 2.3100e+00, ..., 1.5300e+01, 3.9690e+02,
        4.9800e+00],
       [2.7310e-02, 0.0000e+00, 7.0700e+00, ..., 1.7800e+01, 3.9690e+02,
        9.1400e+00],
       [2.7290e-02, 0.0000e+00, 7.0700e+00, ..., 1.7800e+01, 3.9283e+02,
        4.0300e+00],
       ...,
       [6.0760e-02, 0.0000e+00, 1.1930e+01, ..., 2.1000e+01, 3.9690e+02,
        5.6400e+00],
       [1.0959e-01, 0.0000e+00, 1.1930e+01, ..., 2.1000e+01, 3.9345e+02,
        6.4800e+00],
       [4.7410e-02, 0.0000e+00, 1.1930e+01, ..., 2.1000e+01, 3.9690e+02,
        7.8800e+00]])
```

```python
Xbos.std()
```

Output:
```
CRIM         8.601545
ZN          23.322453
INDUS        6.860353
CHAS         0.253994
NOX          0.115878
RM           0.702617
AGE         28.148861
DIS          2.105710
RAD          8.707259
TAX        168.537116
PTRATIO      2.164946
B           91.294864
LSTAT        7.141062
dtype: float64
```

```python
Xbos.mean()
```

Output:
```
CRIM         3.613524
ZN          11.363636
INDUS       11.136779
CHAS         0.069170
NOX          0.554695
RM           6.284634
AGE         68.574901
DIS          3.795043
RAD          9.549407
TAX        408.237154
PTRATIO     18.455534
B          356.674032
LSTAT       12.653063
dtype: float64
```

```python

```

```python

```

```python
type(Xbos)
```

Output:
```
pandas.core.frame.DataFrame
```

Now we are going to manipulate our dataframe Xbos. Keep in mind when querying and changing-- are you copying the     dataframe or altering the actual dataframe, or getting a 'view' of a dataframe without actually altering it.
When reading the Pandas documentation- think about ways to be careful here!
 
Let's start with selecting data with loc and iloc. 
(This is a common interview question for data science jobs/internships, to explain the difference)

This is a pretty decent set of examples to show the difference between loc and iloc:
https://thispointer.com/select-rows-columns-by-name-or-index-in-dataframe-using-loc-iloc-python-pandas/

loc: loc is label-based, which means that we have to specify the name of the rows and columns that we need to filter out.
 
iloc: iloc is integer index-based. So here, we have to specify rows and columns by their integer index.

<b style="color:red">
Starting with loc, try: Xbos.loc where CRIM >= 45. assign this to the variable cri. 
</b> Answer in the cell below.

<!-- grade_id: loc, points: 0, grade: false, solution: true -->
```python
cri = Xbos.loc[Xbos.CRIM >= 45]

cri.info
```

Output:
```
<bound method DataFrame.info of         CRIM   ZN  INDUS  CHAS    NOX     RM    AGE     DIS   RAD    TAX  \
380  88.9762  0.0   18.1   0.0  0.671  6.968   91.9  1.4165  24.0  666.0   
405  67.9208  0.0   18.1   0.0  0.693  5.683  100.0  1.4254  24.0  666.0   
410  51.1358  0.0   18.1   0.0  0.597  5.757  100.0  1.4130  24.0  666.0   
414  45.7461  0.0   18.1   0.0  0.693  4.519  100.0  1.6582  24.0  666.0   
418  73.5341  0.0   18.1   0.0  0.679  5.957  100.0  1.8026  24.0  666.0   

     PTRATIO       B  LSTAT  
380     20.2  396.90  17.21  
405     20.2  384.97  22.98  
410     20.2    2.60  10.11  
414     20.2   88.27  36.98  
418     20.2   16.45  20.62  >
```

<!-- grade_id: loc_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal (cri , Xbos.loc[Xbos.CRIM >= 45])
```

(1 pt.) Now try one with multiple conditions, CRIM >= 45 and  AGE < 100 and assign to  b

<!-- grade_id: multdf, points: 0, grade: false, solution: true -->
```python
b = Xbos.loc[(Xbos.CRIM >= 45) & (Xbos.AGE < 100)]

b.info
```

Output:
```
<bound method DataFrame.info of         CRIM   ZN  INDUS  CHAS    NOX     RM   AGE     DIS   RAD    TAX  \
380  88.9762  0.0   18.1   0.0  0.671  6.968  91.9  1.4165  24.0  666.0   

     PTRATIO      B  LSTAT  
380     20.2  396.9  17.21  >
```

<!-- grade_id: b_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal (b, Xbos.loc[(Xbos.CRIM >= 45) & (Xbos.AGE < 100)] )
```

You can also use loc to update the values of a particular column(s) on selected row(s),
for example you have a dataframe called data and you have rows of information about individuals. You want to select
all people over the age of 18, and then update the column 'age_level' with an 'A' (for adult).
you would do this by

data.loc[(data.age >= 18), ['age_level']] = 'A'
(Hopefully you were wondering, why are the dot method and bracket method used here? Read on for a little explanation here: https://stackoverflow.com/questions/41130255/what-is-the-difference-between-using-squared-brackets-or-dot-to-access-a-column)

(1 pt.) Take your Xbos, find all instances where AGE >80, and update the TAX column to 245.0. Check to see it worked.

<!-- grade_id: loc_update, points: 0, grade: false, solution: true -->
```python
Xbos.loc[(Xbos.AGE > 80), ['TAX']] = 245.0
```

 <b style="color:red">Show your update worked by printing out your previous query below.</b>

```python
print(Xbos.loc[Xbos.AGE > 80])
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM    AGE     DIS  RAD    TAX  \
7    0.14455  12.5   7.87   0.0  0.524  6.172   96.1  5.9505  5.0  245.0   
8    0.21124  12.5   7.87   0.0  0.524  5.631  100.0  6.0821  5.0  245.0   
9    0.17004  12.5   7.87   0.0  0.524  6.004   85.9  6.5921  5.0  245.0   
10   0.22489  12.5   7.87   0.0  0.524  6.377   94.3  6.3467  5.0  245.0   
11   0.11747  12.5   7.87   0.0  0.524  6.009   82.9  6.2267  5.0  245.0   
..       ...   ...    ...   ...    ...    ...    ...     ...  ...    ...   
491  0.10574   0.0  27.74   0.0  0.609  5.983   98.8  1.8681  4.0  245.0   
492  0.11132   0.0  27.74   0.0  0.609  5.983   83.5  2.1099  4.0  245.0   
503  0.06076   0.0  11.93   0.0  0.573  6.976   91.0  2.1675  1.0  245.0   
504  0.10959   0.0  11.93   0.0  0.573  6.794   89.3  2.3889  1.0  245.0   
505  0.04741   0.0  11.93   0.0  0.573  6.030   80.8  2.5050  1.0  245.0   

     PTRATIO       B  LSTAT  
7       15.2  396.90  19.15  
8       15.2  386.63  29.93  
9    
```

Other necessary pandas tasks are to drop columns, copy columns to make new dataframes, and append columns into a table.
To copy a column from Xbos into its own dataframe, you can do this:
newbos = Xbos.TAX.copy()
Then you can drop the TAX column from Xbos:
Xbos.drop(['TAX'], axis =1)
and add it back again.
Xbos.append(newbos)
You will need to do this particularly if your datasets have the target cols included in the dataframe, or separate,
just be aware of this!
On this, you can merge series and dataframes, concatenate, do types of joins, and so on.

Pandas is great for dealing with missing data. You can remove rows or cols with missing data, or impute them to other values such as the mean or mode. 

(1pt.) Run a query that finds NANs in the entire dataset Xbos, both rows and columns. Assign to the variable r and print results.

<!-- grade_id: nans, points: 1, grade: true, solution: true -->
```python
r = Xbos.isnull()
print(r)
```

Output:
```
CRIM     ZN  INDUS   CHAS    NOX     RM    AGE    DIS    RAD    TAX  \
0    False  False  False  False  False  False  False  False  False  False   
1    False  False  False  False  False  False  False  False  False  False   
2    False  False  False  False  False  False  False  False  False  False   
3    False  False  False  False  False  False  False  False  False  False   
4    False  False  False  False  False  False  False  False  False  False   
..     ...    ...    ...    ...    ...    ...    ...    ...    ...    ...   
501  False  False  False  False  False  False  False  False  False  False   
502  False  False  False  False  False  False  False  False  False  False   
503  False  False  False  False  False  False  False  False  False  False   
504  False  False  False  False  False  False  False  False  False  False   
505  False  False  False  False  False  False  False  False  False  False   

     PTRATIO      B  LSTAT  
0      False  False  False  
1      False  False  Fa
```

Let's add a row to Xbos. There are many ways to do this. Create a Pandas series and call it new_row.

```python
new_row = pd.DataFrame({'CRIM':[2], 'ZN':[87], 'INDUS':[92], 'AGE':[97]})
```

Append row to the dataframe

```python
Xbos = pd.concat((Xbos, new_row))
Xbos.tail()
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM   AGE     DIS  RAD    TAX  \
502  0.04527   0.0  11.93   0.0  0.573  6.120  76.7  2.2875  1.0  273.0   
503  0.06076   0.0  11.93   0.0  0.573  6.976  91.0  2.1675  1.0  245.0   
504  0.10959   0.0  11.93   0.0  0.573  6.794  89.3  2.3889  1.0  245.0   
505  0.04741   0.0  11.93   0.0  0.573  6.030  80.8  2.5050  1.0  245.0   
0    2.00000  87.0  92.00   NaN    NaN    NaN  97.0     NaN  NaN    NaN   

     PTRATIO       B  LSTAT  
502     21.0  396.90   9.08  
503     21.0  396.90   5.64  
504     21.0  393.45   6.48  
505     21.0  396.90   7.88  
0        NaN     NaN    NaN
```

Let's add a column of data to Xbos, and give it a default value of 0.

```python
Xbos['Address'] = 0 
Xbos.tail()
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM   AGE     DIS  RAD    TAX  \
502  0.04527   0.0  11.93   0.0  0.573  6.120  76.7  2.2875  1.0  273.0   
503  0.06076   0.0  11.93   0.0  0.573  6.976  91.0  2.1675  1.0  245.0   
504  0.10959   0.0  11.93   0.0  0.573  6.794  89.3  2.3889  1.0  245.0   
505  0.04741   0.0  11.93   0.0  0.573  6.030  80.8  2.5050  1.0  245.0   
0    2.00000  87.0  92.00   NaN    NaN    NaN  97.0     NaN  NaN    NaN   

     PTRATIO       B  LSTAT  Address  
502     21.0  396.90   9.08        0  
503     21.0  396.90   5.64        0  
504     21.0  393.45   6.48        0  
505     21.0  396.90   7.88        0  
0        NaN     NaN    NaN        0
```

## Visualizing and Examining data: the Iris Dataset

In this section, you will start using the scikit-learn Python package and examine a classic dataset used for machine learning. This section will give you practice in thinking about data, which is the often the most important aspect of machine learning. 

Now we are going to import the Iris dataset right from sklearn. There are other versions available as well on the web-- keep that in mind, some vary slightly or are subsets of a larger set! We are going to load the dataset, and examine its properties (you may have to search these pandas commands). We will then look at some preprocessing and visualizations.

```python
from sklearn import datasets # check these datasets out for built ins
# We are going to load the famous iris dataset.
```

```python
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import Axes3D
from sklearn.decomposition import PCA
import numpy as np
import pandas as pd
# import some data to play with
iris = datasets.load_iris()
#type(iris)

data1 = pd.DataFrame(data= np.c_[iris['data'], iris['target']],
                     columns= iris['feature_names'] + ['target'])
# Notice  np.c_  What is this doing? type np.c_? for help.
np.c_?

```

Output:
```
[1;31mType:[0m        CClass
[1;31mString form:[0m <numpy.lib.index_tricks.CClass object at 0x0000011DF5B62B10>
[1;31mLength:[0m      0
[1;31mFile:[0m        c:\users\micha\anaconda3\lib\site-packages\numpy\lib\index_tricks.py
[1;31mDocstring:[0m  
Translates slice objects to concatenation along the second axis.

This is short-hand for ``np.r_['-1,2,0', index expression]``, which is
useful because of its common occurrence. In particular, arrays will be
stacked along their last axis after being upgraded to at least 2-D with
1's post-pended to the shape (column vectors made out of 1-D arrays).

See Also
--------
column_stack : Stack 1-D arrays as columns into a 2-D array.
r_ : For more detailed documentation.

Examples
--------
>>> np.c_[np.array([1,2,3]), np.array([4,5,6])]
array([[1, 4],
       [2, 5],
       [3, 6]])
>>> np.c_[np.array([[1,2,3]]), 0, 0, np.array([[4,5,6]])]
array([[1, 2, 3, ..., 4, 5, 6]])
```

np_c turns sliced objects into concatenated objects.

### Question 12 (6 pts)

In the cells below, look at data1. Answer each of these questions in a separate cell, using numpy and pandas functions.

Use markdown cells to answer the written questions. 

1. General written question: Describe in a few sentences what this data is. You can get it from the actual set, or you can look it up online, either way is fine.
2. What is the data types?
3. What are the dimensions? 
3. Is there any missing data?
4. What is the statistical overview? (mean, std etc.)
5. General written question: What is the target column? What can it be used for?

<b style="color:red">Use the code cells below to get your answers and type in your answers in the markdown cell after that.</b>

```python
print(data1)
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
0                  5.1               3.5                1.4               0.2   
1                  4.9               3.0                1.4               0.2   
2                  4.7               3.2                1.3               0.2   
3                  4.6               3.1                1.5               0.2   
4                  5.0               3.6                1.4               0.2   
..                 ...               ...                ...               ...   
145                6.7               3.0                5.2               2.3   
146                6.3               2.5                5.0               1.9   
147                6.5               3.0                5.2               2.0   
148                6.2               3.4                5.4               2.3   
149                5.9               3.0                5.1               1.8   

     target  
0       0.0  
1   
```

The data contains information on the petal and sepals of various petal irises. The target column is used so the model can characterize the iris from the petal/sepal data.

```python
data1.dtypes
```

Output:
```
sepal length (cm)    float64
sepal width (cm)     float64
petal length (cm)    float64
petal width (cm)     float64
target               float64
dtype: object
```

```python
data1.ndim
```

Output:
```
2
```

```python
data1.isnull().sum().any()
```

Output:
```
False
```

```python
data1.mean()
```

Output:
```
sepal length (cm)    5.843333
sepal width (cm)     3.057333
petal length (cm)    3.758000
petal width (cm)     1.199333
target               1.000000
dtype: float64
```

```python
data1.std()
```

Output:
```
sepal length (cm)    0.828066
sepal width (cm)     0.435866
petal length (cm)    1.765298
petal width (cm)     0.762238
target               0.819232
dtype: float64
```

```python
data1.median()
```

Output:
```
sepal length (cm)    5.80
sepal width (cm)     3.00
petal length (cm)    4.35
petal width (cm)     1.30
target               1.00
dtype: float64
```

```python
data1.max()
```

Output:
```
sepal length (cm)    7.9
sepal width (cm)     4.4
petal length (cm)    6.9
petal width (cm)     2.5
target               2.0
dtype: float64
```

```python
data1.mean()
```

Output:
```
sepal length (cm)    5.843333
sepal width (cm)     3.057333
petal length (cm)    3.758000
petal width (cm)     1.199333
target               1.000000
dtype: float64
```

```python
data1.quantile(0.25)
```

Output:
```
sepal length (cm)    5.1
sepal width (cm)     2.8
petal length (cm)    1.6
petal width (cm)     0.3
target               0.0
Name: 0.25, dtype: float64
```

```python
data1.quantile(0.75)
```

Output:
```
sepal length (cm)    6.4
sepal width (cm)     3.3
petal length (cm)    5.1
petal width (cm)     1.8
target               2.0
Name: 0.75, dtype: float64
```

You *could* rename colums as data1.columns = ['sepal_length', 'sepal_width', 'petal_length', 'petal_width'] What might be wrong with this way of renaming the columns?

By using this naming convention, an outside viewer will not know the units used to describe the data, which may lead to incorrect algorithms that do not take into account the size of each unit. A 6 meter petal is vastly different from a 6 centimeter petal after all.

### Correlation between different features
There's lots to think about when preprocessing a dataset. We'll be going over methods in each of the labs. The first thing you might want to do is look at the correlation between two features. The Pearson correlation coefficient (named for Karl Pearson) can be used to summarize the strength of the linear relationship between two data samples.

The Pearson’s correlation coefficient is calculated as the covariance of the two variables divided by the product of the standard deviation of each data sample. It is the normalization of the covariance between the two variables to give an interpretable score. In statistics, the Pearson correlation coefficient, also referred to as Pearson's r, the Pearson product-moment correlation coefficient (PPMCC), or the bivariate correlation, is a statistic that measures linear correlation between two variables X and Y. It has a value between +1 and −1. A value of +1 is total positive linear correlation, 0 is no linear correlation, and −1 is total negative linear correlation.

In our case, X and Y are two columns of our Iris dataset. While numpy, scipy, and sklearn have methods to calculate correlations of features, we are going to do this "by hand". Use only numpy! Do not use any corr() built in functions!

Bear with me, this may seem pointless or tedious but it is not! The reason is you will be calculating a formula with vectors-- this is definitely something you will see in machine learning and will get you familiar with using numpy.

to calculate r, this is the formula you will translate into numpy:

$$ r = \frac{  n\sum(XY)- (\sum(X)(\sum(Y) }{%
        \sqrt{[n\sum X^{2}-(\sum X)^2][n\sum Y^{2}- (\sum Y)^{2}]}} $$
        
Where X and Y are feature columns (i.e. sepal width and sepal length) and n is the number of elements in each column.


## Question 13 (3 pts)

<b style="color:red">Calculate the Pearson correlations for your Iris dataset in the cell below.</b> (Many ways to do this, if you want to get fancy for practice go for it).  For this correlation question, feel free to use more than one cell, just make it clear.

<!-- grade_id: corrs2, points: 1, grade: true, solution: true -->
```python
def sum_n(arr, power):
    res = 0
    for x in arr: 
        res += x**power
    return res
```

```python
def pearson(arr1, arr2):
    r = 0
    n = data1["sepal length (cm)"].size
    t1 = n * sum_n(arr1 * arr2, 1)
    t2 = sum_n(arr1, 1)
    t3 =sum_n(arr2, 1)
    numerator = t1 - t2 * t3

    t4 = n * sum_n(arr1, 2)
    t5 = n * sum_n(arr2, 2)
    denominator = ((t4 - t2**2) * (t5 - t3**2))**0.5
    return (numerator / denominator)
```

```python
n = data1["sepal length (cm)"].size
pearson(data1["sepal length (cm)"], data1["sepal width (cm)"])
for x in data1: 
    if x == "target": 
        continue
    for y in data1:
        if y == "target":
            continue
        r = pearson(data1[x], data1[y])
        print(f"{x} and {y} have a Pearson Correlation Coefficient of {r}")
```

Output:
```
sepal length (cm) and sepal length (cm) have a Pearson Correlation Coefficient of 1.0
sepal length (cm) and sepal width (cm) have a Pearson Correlation Coefficient of -0.11756978413302598
sepal length (cm) and petal length (cm) have a Pearson Correlation Coefficient of 0.8717537758865967
sepal length (cm) and petal width (cm) have a Pearson Correlation Coefficient of 0.8179411262715885
sepal width (cm) and sepal length (cm) have a Pearson Correlation Coefficient of -0.11756978413302598
sepal width (cm) and sepal width (cm) have a Pearson Correlation Coefficient of 1.0
sepal width (cm) and petal length (cm) have a Pearson Correlation Coefficient of -0.42844010433057067
sepal width (cm) and petal width (cm) have a Pearson Correlation Coefficient of -0.3661259325364594
petal length (cm) and sepal length (cm) have a Pearson Correlation Coefficient of 0.8717537758865967
petal length (cm) and sepal width (cm) have a Pearson Correlation Coefficient of -0.42844010433057067
petal length (cm) an
```

### Question 14 (1 pt)

<b style="color:red">What are the two most correlated features?</b>

To help answer this we can calculate the correlation matrix and create a map of the correlations between features. 

```python
# Here, we show you a seaborn heatmap. It is a matrix of the correlations between features. It uses
# color to illustrate intensity of correlations. This will become very useful when looking at what features 
# are imporant for training your models. 
import seaborn as sns
data2 = data1.drop('target', axis=1)
sns.heatmap(data2.corr(), annot = True);
#annot = True adds the numbers onto the squares
```

Based on your plot, put your answer here

Besides the trivial 1.0 correlations from comparing the feature with itself, the two most correlated features are petal length (measured in cm) and petal width (measured in cm)

### Question 15 (2 pts)

```python
# Let's make a scatter plot of the data

# The indices of the features that we are plotting 
x_index = 0
y_index = 1
# this formatter will label the colorbar with the correct target names
# notice that the target has numerical values, the target_names exist in the iris data
# we just plot right from the dataset, not from our dataframe.
formatter = plt.FuncFormatter(lambda i, *args: iris.target_names[int(i)])

# In the cell below the plot, can you explain what this lambda function is doing?
# Do you understand how the data is gathered and shown?

plt.figure(figsize=(5, 4))
plt.scatter(iris.data[:, x_index], iris.data[:, y_index], c=iris.target)
plt.colorbar(ticks=[0, 1, 2], format=formatter)
plt.xlabel(iris.feature_names[x_index])
plt.ylabel(iris.feature_names[y_index])
plt.tight_layout()
plt.show()
```

(1pt.) <b style="color:red">What is the lambda function/funcformatter doing?</b> Answer in the next cell.


The lambda function takes the numerical target data and assigns it the categorized iris type to format the color bar.

(1pt.) Copying and pasting code, <b style="color:red">now plot petal length vs petal width by changing the x_index and y_index values.</b>

```python
x_index = 2
y_index = 3
# this formatter will label the colorbar with the correct target names
# notice that the target has numerical values, the target_names exist in the iris data
# we just plot right from the dataset, not from our dataframe.
formatter = plt.FuncFormatter(lambda i, *args: iris.target_names[int(i)])

# In the cell below the plot, can you explain what this lambda function is doing?
# Do you understand how the data is gathered and shown?

plt.figure(figsize=(5, 4))
plt.scatter(iris.data[:, x_index], iris.data[:, y_index], c=iris.target)
plt.colorbar(ticks=[0, 1, 2], format=formatter)
plt.xlabel(iris.feature_names[x_index])
plt.ylabel(iris.feature_names[y_index])
plt.tight_layout()
plt.show()
```

Pandas built-in Viz
Pandas has built in data visualization, based on matplotlib, so we show a few examples. (Feel free to use any plotting package you like for this course, but these are pretty convenient if you have a pandas dataframe ready to go.

```python
data1['sepal length (cm)'].hist()
```

Output:
```
<Axes: xlabel='petal length (cm)', ylabel='petal width (cm)'>
```

```python
data1.plot.scatter(x='sepal length (cm)', y= 'sepal width (cm)', s=data1['target']*200)
```

Output:
```
<Axes: xlabel='sepal length (cm)', ylabel='sepal width (cm)'>
```

```python
# Let's do some plotting with seaborn.
# We start with a jointplot
sns.set_style('whitegrid')
sns.jointplot(x= 'petal width (cm)', y = 'petal length (cm)', data=data1)
```

Output:
```
<seaborn.axisgrid.JointGrid at 0x11d8183a550>
```

<!-- grade_id: distplot, points: 1, grade: true, solution: true -->
```python
# (1pt.) Using data1 and petal width, create a seaborn displot with 30 bins, kde=True, color='red'

### BEGIN SOLUTION
sns.displot(data = data1, x = "petal width (cm)", bins = 30, kde = True, color = 'red')
### END SOLUTION
```

Output:
```
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\axisgrid.py:118: UserWarning: The figure layout has changed to tight
  self._figure.tight_layout(*args, **kwargs)
<seaborn.axisgrid.FacetGrid at 0x11d817dd450>
```

## Plotting Data
For this final section of the lab, we are going to be plotting using seaborn and the Titanic dataset. Your tasks will be to plot various directives. (There are also lots of Jupyter notebooks on Kaggle.com from challenges using the Titanic dataset, FYI)

```python
titanic = sns.load_dataset('titanic') #notice where we are getting Titanic from!
titanic.head()
```

Output:
```
survived  pclass     sex   age  sibsp  parch     fare embarked  class  \
0         0       3    male  22.0      1      0   7.2500        S  Third   
1         1       1  female  38.0      1      0  71.2833        C  First   
2         1       3  female  26.0      0      0   7.9250        S  Third   
3         1       1  female  35.0      1      0  53.1000        S  First   
4         0       3    male  35.0      0      0   8.0500        S  Third   

     who  adult_male deck  embark_town alive  alone  
0    man        True  NaN  Southampton    no  False  
1  woman       False    C    Cherbourg   yes  False  
2  woman       False  NaN  Southampton   yes   True  
3  woman       False    C  Southampton   yes  False  
4    man        True  NaN  Southampton    no   True
```

### Question 16 (16 pts)
For this question you are going to make a series of plots with seaborn in the cells below using the Titanic dataset. **Remember to label your axes!** Include units if you have them. 

Part 1: In a cell below, create a correlation matrix of features (you can use any method you like)


Part 2: Using seaborn and Titanic, in a cell below create a joint plot with the x axis as fare, the y axis as age. feel free to experiment with size and/or colors if you like.

Part 3: In a cell below, create a histogram of the fare, kde= False, and choose an appropriate number of bins


Part 4: Create a box plot with x axis being class type, and y axis age. Try palette ='rainbow'

Part 5: Create a swarm plot with x axis class, and y axis age, try the palette 'Set2'.
You might get a warning-- can you fix it? (it is okay to leave the warning, however)

Part 6: Create a count plot with x axis using sex.


Part 7: Create a heatmap of the correlations. Title it "Titanic"
 

Part 8: Create a facet grid using sex as the columns, then map the histogram on age.


```python
# answer people are missing
# g = sns.FacetGrid(data=titanic,col='sex')
# g.map(plt.hist,'age')
```

```python
#Part 1
features = titanic.select_dtypes('number')
sns.heatmap(features.corr(), annot = True)
```

Output:
```
<Axes: >
```

```python
#Part 2
sns.set_style("white")
sns.jointplot(x = "fare", y = "age", data = titanic, color = 'green') 
```

Output:
```
<seaborn.axisgrid.JointGrid at 0x11d89df80d0>
```

```python
#Part 3
titanic['fare'].hist(bins = 50)
```

Output:
```
<Axes: xlabel='fare', ylabel='age'>
```

```python
#Part 4
sns.boxplot(data = titanic, x = "class", y = "age", palette = "rainbow", hue = "class")
```

Output:
```
<Axes: xlabel='class', ylabel='age'>
```

```python
#Part 5
sns.swarmplot(data = titanic, x = "class", y = "age", palette = "Set2", hue = "class", size = 3)
```

Output:
```
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\categorical.py:3544: UserWarning: 94.6% of the points cannot be placed; you may want to decrease the size of the markers or use stripplot.
  warnings.warn(msg, UserWarning)
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\categorical.py:3544: UserWarning: 95.4% of the points cannot be placed; you may want to decrease the size of the markers or use stripplot.
  warnings.warn(msg, UserWarning)
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\categorical.py:3544: UserWarning: 95.5% of the points cannot be placed; you may want to decrease the size of the markers or use stripplot.
  warnings.warn(msg, UserWarning)
<Axes: xlabel='class', ylabel='age'>
```

```python
#Part 6
sns.countplot(data = titanic, x = "sex")
```

Output:
```
<Axes: xlabel='sex', ylabel='count'>
```

```python
#Part 7
features = titanic.select_dtypes('number')
ax = plt.axes()
sns.heatmap(features.corr(), annot = True, ax = ax)
ax.set_title("Titanic")
```

Output:
```
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\categorical.py:3544: UserWarning: 23.4% of the points cannot be placed; you may want to decrease the size of the markers or use stripplot.
  warnings.warn(msg, UserWarning)
Text(0.5, 1.0, 'Titanic')
```

```python
#Part 7
grid = sns.FacetGrid(data = titanic, col = "sex")
grid.map(titanic['age'].hist)
```

Output:
```
c:\Users\micha\anaconda3\Lib\site-packages\seaborn\axisgrid.py:118: UserWarning: The figure layout has changed to tight
  self._figure.tight_layout(*args, **kwargs)
<seaborn.axisgrid.FacetGrid at 0x11d8a3b12d0>
```

# Good Job!
Rerun your cells and submit your assignment. 